Human receptive endometrial assembloid for deciphering the implantation window

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

This study generated 3D cell constructs from endometrial cell mixtures that were seeded in the Matrigel scaffold. The cell assemblies were treated with hormones to induce a "window of implantation" (WOI) state. Although many bioinformatic analyses point in this direction, there are major concerns that must be addressed.

Strengths:

The addition of 3 hormones to enhance the WOI state (although not clearly supported in comparison to the secretory state).

Comments on revisions:

The authors did their best to revise their study according to the Reviewers' comments. However, the study remains unconvincing, incomplete and at the same time still too dense and not focused enough.

Reviewer #2 (Public review):

Zhang et al. have developed an advanced three-dimensional culture system of human endometrial cells, termed a receptive endometrial assembloid, that models the uterine lining during the crucial window of implantation (WOI). During this mid-secretory phase of the menstrual cycle, the endometrium becomes receptive to an embryo, undergoing distinctive changes. In this work, endometrial cells (epithelial glands, stromal cells, and immune cells from patient samples) were grown into spheroid assembloids and treated with a sequence of hormones to mimic the natural cycle. Notably, the authors added pregnancy-related factors (such as hCG and placental lactogen) on top of estrogen and progesterone, pushing the tissue construct into a highly differentiated, receptive state. The resulting WOI assembloid closely resembles a natural receptive endometrium in both structure and function. The cultures form characteristic surface structures like pinopodes and exhibit abundant motile cilia on the epithelial cells, both known hallmarks of the mid-secretory phase. The assembloids also show signs of stromal cell decidualization and an epithelial mesenchymal transition, like process at the implantation interface, reflecting how real endometrial cells prepare for possible embryo invasion.

Although the WOI assembloid represents an important step forward, it still has limitations: the supportive stromal and immune cell populations decrease over time in culture, so only earlypassage assembloids retain full complexity. Additionally, the differences between the WOI assembloid and a conventional secretory-phase organoid are more quantitative than absolute; both respond to hormones and develop secretory features, but the WOI assembloid achieves a higher degree of differentiation due to the addition of "pregnancy" signals. Overall, while it's a reinforced model (not an exact replica of the natural endometrium), it provides a valuable in vitro system for implantation studies and testing potential interventions, with opportunities to improve its long-term stability and biological fidelity in the future.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

This study generated 3D cell constructs (i.e., assembloids) that were treated with hormones to induce a 'window of implantation' (WOI) state. While the authors have made large efforts to address the reviewers' feedback, the study's findings remain unconvincing and incomplete.

(1) The authors have appropriately revised the terminology from 'organoids' to 'assembloids' in several parts of the manuscript. However, this revision remains incomplete, as the main title, figure legends, and figure titles still contain the incorrect term. A thorough review of the entire manuscript is recommended to ensure consistent and accurate use of terminology.

Thank you for your meticulous review. We have now conducted a full check and confirmed that terminology is used consistently and accurately throughout the text.

(1) Previous comments raised concerns about the feasibility of robustly passaging assembloid structures - comprising epithelial, stromal and immune cells - under epithelial growth conditions. The authors responded by stating that they optimized the expansion medium with a stromal cell-promoting factor. Additionally, rather than conducting scRNA-seq on both early and late passages (P6-P10) as suggested, they performed immunofluorescence staining, which confirmed the persistence of stromal cells at passage 6. However, the presence of immune cells was not addressed. Confirmation of their presence is essential for all further claims. Moreover, a more zoomed-out view of the immunostaining would help clarify the overall cellular composition across the entire well and facilitate comparison with corresponding brightfield images.

Whole-mount immunofluorescence of the 6th - generation assembloids revealed that CD45+ immune cells surrounded FOXA2+ glands, with a more zoomed-out view provided.

Whole-mount immunofluorescence showed that CD45+ cells (immune cells) were arranged around the glandular spheres that were FOXA2+.

Scale bar =50 μm (left) and 30 μm (right).

In their response, the authors mention using the first three passages to ensure optimal cell diversity and viability. However, the manuscript states that 'assembloids derived from the first generation are used for experiments' (line 106). This discrepancy must be clarified.

Thank you for your suggestion. We have revised the relevant content to “The assembloids derived from the first three generation are used for experiments” (Line 90-91).

(2) The authors have made a commendable effort to bring more focus to the manuscript, which has improved readability.

We thank you for your insightful suggestions, which have greatly improved the quality of our manuscript.

(3) The "embryo implantation" part remains very unconvincing. How did authors define "the blastoids could grow within the endometrial assembloids and interact with them"? What did they mean with "grow"? Did blastoids further differentiate? Normally, blastoids cannot further "grow". "Survival rates of blastoids" is not equal to "growth". It is not clear how the survival rate was quantified. Besides, regarding the "interaction rates", how did authors define and quantify it? Actually, blastoids are able to attach to Matrigel efficiently (even without any endometrial cells), so authors cannot simply define the "interaction" as the co-localization of blastoids and assembloids via brightfield images. In addition, for the assembloids as the 3D structures grow in the Matrigel, the epithelial parts are normally apical-in, while the blastoids attach to the apical (lumen) side of the epithelial cells, so physiologically, blastoids should interact with the apical part of the epithelial cells instead of the outside of the assembloids.

(1) What did they mean with "grow"? Did blastoids further differentiate?

On the one hand, volume and morphology undergo continuous dynamic changes; on the other hand, only the inner cell mass and trophectoderm exist at the blastocyst stage, with the ICM further differentiating into OCT4+ epiblast and GATA6+ hypoblast.

(2) Survival rates of blastoids" is not equal to "growth". It is not clear how the survival rate was quantified.

The definition of "survival rate" is as follows: morphologically, the blastocoel remains noncollapsed and the cell boundaries are distinct (with no obvious cell detachment); molecularly, the markers of epiblast, hypoblast and trophectoderm are expressed. The survival rate is calculated as the ratio of viable embryoids to the total number of embryoids.

(3) Besides, regarding the "interaction rates", how did authors define and quantify it? Actually, blastoids are able to attach to Matrigel efficiently (even without any endometrial cells), so authors cannot simply define the "interaction" as the co-localization of blastoids and assembloids via brightfield images.

The criteria for determining interaction include not only attachment between the blastoids and assembloids observed via brightfield images, but also their sustained tight adhesion against external mechanical perturbations (e.g., medium replacement, immunostaining procedures).

(4) In addition, for the assembloids as the 3D structures grow in the Matrigel, the epithelial parts are normally apical-in, while the blastoids attach to the apical (lumen) side of the epithelial cells, so physiologically, blastoids should interact with the apical part of the epithelial cells instead of the outside of the assembloids.

You are absolutely correct. In vivo, the embryo indeed makes initial contact with the apical side of the epithelial cells. The introduction of the blastoid co-culture model herein is intended to demonstrate that this receptive endometrial assembloids can better support blastoid growth and development.

(4) Previous comments highlighted the absence of distinct shifts in gene expression profiles between SEC assembloids and WOI assembloids, which contrasts with findings from primary endometrial tissue reported by Wang et al. (2020). While the authors have expanded their analysis using the Mfuzz algorithm and identified changes in mitochondria- and cilia-associated genes, the manuscript still lacks evidence of significant transcriptional changes in key WOI marker genes, as described in Wang et al. This discrepancy must be addressed and discussed in greater depth to clarify the biological relevance of their model.

The endometrium in vivo involves complex crosstalk among multiple cell types and is tightly regulated by the hypothalamic-pituitary-ovarian (HPO) axis, thus exhibiting distinct shifts in gene expression during the peri-implantation period.

In our in vitro model, alterations in mitochondria- and cilia-related genes were observed, which to a certain extent demonstrates that these window of implantation (WOI) assembloids possess receptive-phase characteristics and can be employed to investigate WOI-associated scientific questions or conduct in vitro drug screening.

However, substantial efforts are still required to optimize the current model for fully recapitulating the dynamic changes in endometrial gene expression across different phases in vivo, and this aspect is further addressed in the Limitations section of our discussion (Line 342-353).

“However, our WOI endometrial assembloids also exhibit some limitations. It is undeniable that the assembloids cannot perfectly replicate the in vivo endometrium, which comprises functional and basal layers with a greater abundance of cell subtypes, under superior regulation by hypothalamic-pituitary-ovarian (HPO) axis. Specifically, stromal and immune cells are challenging to stably passage, and their proportion is lower than in the in vivo endometrium. While the in vivo peri-implantation period exhibits intricate gene expression dynamics driven by systemic regulation, our models only partially recapitulate these changes, primarily in mitochondria- and cilia-associated genes. Nevertheless, to some extent, these WOI assembloids possess receptivity characteristics and can be utilized for investigating receptivity-related scientific questions or conducting in vitro drug screening. Further refinements are required to fully simulate the dynamic endometrial gene expression patterns across all menstrual cycle stages. We are looking forward to integrating stem cell induction, 3D printing, and microfluidic systems to modify the culture environment.”

(5) In the authors' response document, they present data integrating their results with those of Garcia Alonso et al. (2021). However, these integrated analyses are not included in the revised manuscript (which should be, if answering a major concern).

Thanks for your valuable suggestions. We have now integrated the findings of Garcia Alonso et al. (2021) into the revised manuscript (Line 132) and Figure S2E–F.

(8) Fig 2D: The authors have clarified that CD45+ staining is used. However, they have not yet adapted the typo in the figure legend of the right picture.

Thanks for your thorough review. The left panel of Figure 2D is stained with CD45 to label immune cells, while the right panel is stained with CD44. These details have been clearly indicated in both the manuscript and the figure legend.

(9) All quantification analyses (as described in the authors' response document) should be clearly described in the Materials & Methods section.

Thanks for your valuable suggestions. All quantification analyses have now been added to the Supporting Materials and Methods section (Line 94-104, Line 110-111, Line 241244).

(10) The authors have provided clarification regarding their method for quantifying immunofluorescence staining (e.g., OLFM4 expression in Fig. 3C) in their response document. However, these methodological details are not included in the revised manuscript. It is important that such information is incorporated into the manuscript itself to ensure transparency and reproducibility for others.

Thanks for your valuable suggestions. All quantification analyses have now been added to the Supporting Materials and Methods section (Line 94-104).

(13) It is needed to include the author's response to the comment about literature showing the opposite of increased number of cilia during the WOI into the discussion part of the paper.

We appreciate your suggestions. The relevant content has now been added to the Discussion section (Lines 319–323).

(14) In the authors' response, they explain the difference between pinopodes and microvilli. They should include this explanation briefly in the manuscript. Moreover, Fig. 3F lacks a picture of cilia structure in CTRL condition. In addition, the structures that are indicated as cilia with an orange arrow seem to not be attached to the endometrial cells (anymore). It would be useful to show another more representative picture for the cilia.

(1) Thank you for your valuable suggestions. The distinction between pinopodes and microvilli has now been added to the Supporting Materials and Methods section (Line 230-236).

(2) You are probably referring to Figure 2F—we did not observe ciliary structures in the CTRL group.

(3) The cilia structure was visualized via transmission electron microscopy (TEM), which requires ultrathin sectioning. Thus, the cilia shown in the image correspond to a single cross-section of the captured assembloids. Owing to technical limitations, three-dimensional visualization of cilia on the cells cannot be achieved.

(17) The results on co-culturing blastoids with the WOI assembloids is not convincing. The blastoids are exposed to the basolateral side of the endometrial epithelial cells, while in vivo, blastocysts interact with the apical side of the endometrial epithelial cells first (apposition and attachment), followed by invasion into the endometrium. This means that the interaction shown here is not physiological. Therefore, it is not justified to say that this platform holds promise to investigate maternal-fetal interactions.

We agree with your perspective that discrepancies exist between this model and the physiological processes in vivo. However, such differences do not negate the scientific value of the model.

The core merit of this study lies in the successful establishment of co-culture systems for blastoids and WOI assembloids. Notably, genuine cross-talk occurs between the two components, thereby providing a practical and operational tool for subsequent research.

Although the current contact orientation differs from that observed in vivo, future optimization of the cell culture protocol (via modulation of cell polarity) will enable the model to better recapitulate physiological conditions. Therefore, the innovation and operability of this model within specific research contexts still render it a robust platform for investigating maternal-fetal interactions.

Overall, it is highly recommended that the authors carefully review the manuscript for grammatical errors, inconsistencies and issues with scientific phrasing. The language throughout the text requires substantial editing to improve clarity, readability and precision.

We appreciate your suggestions. A full manuscript check was performed to rectify grammatical errors, inconsistencies, and inappropriate scientific phrasing, with further language refinement by a native English-speaking specialist.

Fig 1A: This overview is unclear. How many days do the assembloids grow before being stimulated with hormones? Are CTRL assembloids only kept in culture until day 2 and SEC and WOI assembloids until day 8? This is also not clear form the Materials and Methods section. Should be clarified.

Thanks for your valuable suggestions. We have now updated the overview (Figure 1A) and Materials and Methods section (Line 370-371, Line 379-381).

“Hormonal treatment was initiated following the assembly of the endometrial assembloids (about 7-day growth period).”

“The CTRL group was cultured in ExM without hormone supplementation and subjected to parallel culture for 8 days along with the two aforementioned groups.”

Fig 1B: From these brightfield images, it appears that the size of the assembloids remains relatively consistent from Day 0 to Day 3 and up to Day 11 (especially in CTRL). However, in Fig S1A, the assembloids on Day 11 appear significantly larger compared to those on Day 2 (or Day 4). Authors should clarify this discrepancy (since both of the figures are shown as "brightfield of endometrial assembloids").

You are probably referring to the observation that the assembloids at Day 11 in Fig. S1A are smaller in size than those at Day 2 (or Day 4) in Fig. 1B. This discrepancy arises because the time points in Fig. 1B are calculated starting from the initiation of hormone treatment for the SEC and WOI groups, rather than from the beginning of the overall culture as in Fig. S1A. In addition, assembloids exhibit size variability during the same culture period due to individual heterogeneity.

To eliminate ambiguity, we have now labeled “Hormone Day 0, Day 2, Day 8” in Fig. 1B and revised the corresponding figure legend to read: “Endometrial assembloids from the CTRL, SEC, and WOI groups, which were subjected to hormone treatment on Days 0, 2, and 8, exhibited comparable growth patterns throughout the culture period.”

Fig 2G: authors still used the description "organoids" here instead of "assembloids".

We appreciate your careful review. Corrections have been made accordingly.

Fig. 3C: For the OLFM4 staining quantification, in the Y-axis authors wrote "proportion of OLFM4 (+) cells (OLFM4 (+)/total", but in the rebuttal letter they mention "its fluorescence intensity (quantified as mean grey value) was significantly stronger in both the SEC and WOI groups compared to the CTRL group". This is confounding and should be clarified.

We apologize for incorrectly writing "fluorescence intensity" in the rebuttal letter; the correct term should be the "proportion of OLFM4 (+) cells (OLFM4 (+)/total)" as shown in Fig. 3C.

Fig 5D: Acetyl-α-tubulin is the marker of ciliated cells and should be expressed in the cilia instead of the whole cells. It is very strange to quantify as "mean fluorescence intensity (acetyl-αtubulin/DAPI)" to assess the cilia. Please clarify.

Thank you for your insightful comment. To clarify, the ratio "mean fluorescence intensity (acetyl-α-tubulin/DAPI)" was calculated within individual acetyl-α-tubulin+ ciliated cells. Acetyl-αtubulin fluorescence was normalized to the DAPI signal of the same cell nucleus, not the wholecell population. This corrected for variations in cell number and staining efficiency to ensure data accuracy.

Fig 5F: it is very bizarre that unciliated epithelium was transformed from ciliated epithelium, and CTRL was transformed from SEC and WOI. Should be clarified and discussed.

Pseudotime analysis sorts discrete cells along a "pseudotime axis" based on similarities and differences in cellular gene expression, thereby simulating cell state transitions.

Ciliated epithelium → unciliated epithelium: During the menstrual cycle, ciliated and unciliated epithelia undergo mutual transformation from the secretory phase (or mid-secretory phase) to the menstrual phase, and then to the proliferative phase. Here, we demonstrate the transition of ciliated cells to unciliated cells from the SEC and WOI stages to the CTRL stage.

Notably, the two cell types coexist, and what is presented here merely reflects a transformation trend. Relative content has been incorporated into the Discussion section (Line 319-321).

“Throughout the menstrual cycle, ciliated and unciliated epithelia undergo mutual transformation from the secretory phase (or mid-secretory phase) to the menstrual phase, and then to the proliferative phase.”

Fig 5H: To show "enhanced invasion ability", authors must provide some quantification and statistic analysis. It is very hard to see the difference between the CTRL and SEC regarding ROR2Wnt5A.

We appreciate your suggestion. Quantification and statistic analysis have been added to Figure 5H.

Fig 6A: please elaborate the "mIVC1" and "mIVC2" in the figure legends.

Additions have been made to the figure legends accordingly, as follows: "mIVC1: modified In Vitro Culture Medium 1; mIVC2: modified In Vitro Culture Medium 2."

Fig S1D: Is the PAS staining also done in CTRL assembloids? In addition, it is stated that the assembloids secrete glycogen because of a positive PAS staining, while it could also be neutral mucins, glycoproteins, etc, which are all detected by PAS staining. So, the authors should be more careful in stating that it is glycogen, or a PAS staining with diastase digestion should be done.

The PAS staining results for the CTRL group are presented in Fig. S1I. In addition, results of PAS staining with diastase digestion are included in Figure S1.

Line 120: references?

The reference has been added accordingly.

Line 178: The term 'Endometrial Receptivity Test (ERT)' is used. Do the authors mean Endometrial Receptivity Analysis (ERA) test? ERA is the commonly used abbreviation for this test. Moreover, the authors describe ERA as 'a kind of gene analysis-based test.' This should be rephrased more scientifically correct.

Thank you for your valuable suggestion. We have revised the term to ERA, and modified the phrase "a kind of gene analysis-based test" to "gene expression profiling-based diagnostic assay" (Lines 160–163).

“We performed Endometrial Receptivity Analysis (ERA), a gene expression profiling-based diagnostic assay that integrates high-throughput sequencing and machine learning to quantify the expression of endometrial receptivity-associated genes.”

Line 83: assemblies à assembloids

We appreciate your suggestion. The text has been updated to “the endometrial assembloids progressed from epithelial organoids, to assemblies of epithelial and stromal cells and then to stem cell-laden 3D artificial endometrium”.

The Materials and Methods section currently lacks the needed details. Authors should substantially expand this section to clearly describe all experimental and analytical procedures, including, aùmong others, immunofluorescence staining, quantification methods, bioinformatics analyses and statistical approaches. Providing comprehensive methodological information is essential.

A detailed description of these methods is provided in the Supporting Materials and Methods section.

Reviewer #2 (Recommendations for the authors):

The revised manuscript is much improved in clarity, focus, and experimental support. The authors have thoughtfully addressed the major concerns from the previous review. In particular, the logic and flow of the paper are clearer, it now guides the reader through the rationale (constructing a WOI model), the comparative analysis against in vivo tissue and simpler organoids, and the key features that distinguish the WOI assembloid. The added functional validation (especially the blastoid co-culture experiment) significantly strengthens the work by showing a tangible outcome of "receptivity" beyond molecular profiling. The distinction between the standard secretory-phase organoid and the WOI assembloid is now more convincing, as the authors highlight several specific differences in morphology (more cilia, pinopodes), metabolism, and implantation success that favor the WOI model. The manuscript also reads cleaner with the bioinformatic sections condensed to the most important findings (excess detail was trimmed or moved to supplements) and the rationale for gene/pathway selection explicitly stated.

The manuscript has been significantly strengthened through the addition of functional assays (like the blastoid co-culture), clearer transcriptomic and proteomic data, and detailed analyses of hormone treatments, cilia biology, and stromal and immune cell behavior in early passages. These updates confirm that the WOI assembloid supports embryo attachment and outperforms standard secretory organoids, while integrating external references and clarifications on terminology. Minor suggestions remain, such as clarifying statistical significance and adding functional interpretations for certain observations, but overall, the manuscript is now more robust and biologically convincing.

Remaining points for clarification: There are a few minor points that still merit attention:

- Use of the Endometrial Receptivity Test (ERT): As previously mentioned, if the authors have ERT data for the SEC organoid group, including that information would further support the claim that the WOI assembloid is uniquely receptive. If not, it would be helpful to add a statement clarifying that the ERT was employed specifically as a confirmatory test for the WOI assembloids, rather than as a comparative measure across all groups.

Thank you for your valuable suggestion. We have now supplemented the description in the Supporting Materials and Methods section (Lines 160–162) as follows: “ERA was employed specifically as a confirmatory test for the WOI assembloids, rather than as a comparative measure across all groups.”

- Because the assembloids are created from primary tissue samples, it would be helpful to briefly comment on how consistent the findings were across different patient-derived samples. For example, did all biological replicates show similar expression of receptivity markers and comparable capacity to support blastoid attachment? Although this seems implied, including a sentence in the Methods or Results sections that specifies the number of donor lines tested would help readers assess the model's variability and reproducibility.

We appreciated your advice. The relevant statement has been added to the Supporting Materials and Methods section. (Line 312-313).

“All biological replicates (fourteen individuals) of endometrial assembloids show similar expression of receptivity markers and comparable capacity to support blastoid attachment.”

- The authors mention promising future directions, such as integrating 3D printing and microfluidics to further enhance the model, which is an excellent forward-looking statement. It would also be valuable to suggest the inclusion of additional cell types, like more robust immune cell populations or endothelial components, as future improvements to create an even more comprehensive model of the endometrial lining.

Thank you for your valuable suggestion. 3D printing and microfluidics serve as approaches for introducing multiple cell types. We have supplemented the following statement in the manuscript: “We are looking forward to integrating stem cell induction, 3D printing, and microfluidic systems to modify the culture environment.” (Line 352-353).

We are grateful for your valuable feedback and constructive criticism, which have helped us improve the quality of our work in terms of content and presentation. We have diligently revised the manuscript and made necessary changes. Here, we have attached the revised manuscript, figures, and all supplementary materials for your re-evaluation. Thank you again for your continued support and look forward to your favorable decision.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Human endometrial receptivity is a critical determinant of pregnancy success; however, in vivo studies of its features and regulation are particularly challenging due to ethical restriction. Recently, the development of human endometrial assembloids has provided a powerful model to examine this intricate biological process. In this study, we established a specialized human window-of-implantation (WOI) endometrial assembloid system that mimics the in vivo receptive endometrium. It recapitulates not only the structural attributes of pinopodes and cilia, but also molecular characteristics of mid secretory endometrium. Furthermore, the WOI endometrial assembloid exhibits hormone responsiveness, energy metabolism with larger and enhanced functional mitochondria, increased ciliary assembly and motility, and epithelial-mesenchymal transition (EMT), as well as promising potential for embryo implantation. As such, WOI assembloids hold great promise as a platform to unravel the intricate mechanisms governing endometrial receptivity regulation, maternal-fetal interactions, and associated diseases, ultimately driving impactful advancements in the field.

Revised version: Human endometrial receptivity is a critical determinant of pregnancy success; however, in vivo studies of its features and regulation are particularly challenging due to ethical restriction. Recently, the development of human endometrial assembloids has provided a powerful model to investigate this intricate biological process. In this study, we established a specialized human window-of-implantation (WOI) endometrial assembloid system that mimics the in vivo receptive endometrium. It not only reproduces the structural attributes of pinopodes and cilia, but also molecular characteristics of mid-secretory endometrium. Furthermore, the WOI endometrial assembloid exhibits hormone responsiveness, an energy metabolism profile characterized by larger and functionally enhanced mitochondria, increased ciliary assembly and motility, and epithelial-mesenchymal transition (EMT), as well as promising potential for embryo implantation. As such, WOI assembloids hold great promise as a platform to unravel the intricate mechanisms governing the regulation of endometrial receptivity, maternal-fetal interactions, and associated pathologies, ultimately driving impactful advancements in the field.

Introduction

Earlier version: The human endometrium, comprising various cell types, experiences shedding (menstrual phase), regeneration (proliferative phase), and differentiation (secretory phase) under the regulation of estrogen and progesterone during the menstrual cycle (1). A brief window, called the “window-of-implantation (WOI)” or mid-secretory phase, allows for embryo implantation into the endometrium (1). The endometrium at the embryo implantation site contains diverse cells, including ciliated epithelial cells, decidualized stromal cells and immune cells. Cyclic endometrial changes akin to those in humans are exclusive to apes, Old World monkeys, molossid bats, and spiny mice, but not mice (1), rendering typical mouse models insufficient for accurately simulating human endometrium. Existing endometrial cell lines, such as Ishikawa and HEEC cells, consist solely of a single cell type, thus failing to reproduce the intricate physiological structure and function of the endometrium.

Revised version: The human endometrium, a complex tissue comprising diverse cell types, undergoes shedding (menstrual phase), regeneration (proliferative phase), and differentiation (secretory phase) under the coordinated regulation of estrogen and progesterone throughout the menstrual cycle (1). A brief interval, termed the “window-of-implantation (WOI)” or mid-secretory phase, peimits embryo implantation into the endometrium (1). The endometrium at the site of embryo implantation harbors a heterogeneous population of cells, including ciliated epithelial cells, decidualized stromal cells and immune cells. Cyclic endometrial changes analogous to those observed in humans are exclusive to apes, Old World monkeys, molossid bats, and spiny mice, but are absent in conventional laboratory mice (1). This renders typical mouse models inadequate for accurately recapitulating human endometrium. Existing endometrial cell lines, such as Ishikawa and HEEC cells, are composed of a single cell type, and thus fail to reproduce the intricate physiological architecture and function of the endometrium.

Earlier version: Replicating and reconstructing human organs has become essential for exploring tissue physiology and function. O rganoid, a self-assembled 3D structure, closely resemble in vivo tissue or organ (2). They offer high expansibility, phenotypic, and functional properties, emerging as powerful tools for investigating tissue physiology and disease (3). In 2017, the first long-term and hormone-responsive human endometrial organoid was established using adult stem cells from endometrial biopsy(2, 3). Based on this, Margherita Y. Turco, et al developed noninvasive methods to construct endometrial organoids from menstrual flow(4). Takahiro Arima et al. developed polar reversal endometrial organoids to study embryo implantation(5). Apart from adult stem cells, pluripotent stem cells were also induced to endometrial stromal fibroblasts and epithelium, and then cocultured to form organoids(6, 7), which offered vigorous proliferative capacity but lacked immune cells and other components of the microenvironment. There are also studies that add immune cells to endometrial organoids for co-culture(8). In efforts to mimic human endometrial cell types in vitro, the endometrial organoids progressed from epithelial organoids(9), to assemblies of epithelial and stromal cells(10, 11) and then to stem cell-laden 3D artificial endometrium(12, 13), which were solely closer but not completely identical to the endometrium. Additionally, pathological endometrial organoid models have been established for conditions like endometriosis, endometrial hyperplasia, Lynch syndrome, and endometrial cancer (14). These organoids simulated endometrial morphology, hormone responsiveness, and physiological and pathological processes in vitro (3, 9, 15–17), facilitating the study of physiological phenomena (15, 18, 19), pathogenic mechanisms(16) and drug screening(14).

Revised version: Replicating and reconstructing human organs has become indispensable for elucidating tissue physiology and function. Organoid, a self-assembled 3D structure, closely resemble in vivo tissue or organ (2). They offer high expansibility, alongside conserved phenotypic, and functional properties, emerging as powerful tools for investigating tissue physiology and disease (3). In 2017, the first long-term and hormone-responsive human endometrial organoid was derived from adult stem cells in endometrial biopsies (2, 3). Based on this, Margherita Y. Turco, et al generated endometrial organoids from menstrual effluent using non-invasive methods (4), while Takahiro Arima et al. engineered polarity-reversed organoids to study embryo implantation(5). Beyond adult stem cells, pluripotent stem cells were also induced to endometrial stromal fibroblasts and epithelium, and then cocultured to form organoids(6, 7), which exhibited vigorous proliferative capacity but lacked immune cells and other components of the microenvironment. Several studies have additionally incorporated immune cells into endometrial organoid co-culture systems(8). To recapitulate the full spectrum of human endometrial cell types in vitro, endometrial assembloids have evolved from epithelial organoids(9), to assemblies of epithelial and stromal cells(10, 11) and then to stem cell-laden 3D artificial endometrium(12, 13), which were solely closer but not completely identical to the endometrium. Additionally, pathological endometrial organoid models have been established for conditions like endometriosis, endometrial hyperplasia, Lynch syndrome, and endometrial cancer (14). These organoids faithfully simulated endometrial morphology, hormone responsiveness, and physiological and pathological processes in vitro (3, 9, 15–17), facilitating the study of physiological phenomena (15, 18, 19), pathogenic mechanisms(16) and drug screening(14).

Earlier version: Although various regulators have been suggested to be involved in WOI, we are far from understanding how embryo implantation occurs during the WOI due to ethical limitations and fewer in vitro receptive endometrial model. Transforming from proliferative to receptive endometrium involves dynamic changes like decidualization, epithelial-mesenchymal transition (EMT) and ciliated epithelial development (1, 20), which have not been reconstructed in most in vitro model to date. Here, we successfully established a receptive endometrial assembloid system regulated by hormones and reflecting the in vivo WOI endometrial features, providing a platform for studying physiological and pathological endometrium and maternal-fetal interactions.

Revised version: Although various regulators are implicated in the WOI, our understanding of embryo implantation during this period remains limited by ethical constraints and a paucity of in vitro receptive endometrial models. The transition from proliferative to receptive endometrium involves dynamic processes, including decidualization, epithelial-mesenchymal transition (EMT) and ciliated epithelial development (1, 20), which most in vitro models have yet to fully recapitulate. Here, we established a hormone-regulated receptive endometrial assembloid system that recapitulates in vivo WOI features, providing a platform to study physiological/pathological endometrial function and maternal–fetal interactions.

Results

Developing receptive endometrial assembloids in vitro

Earlier version: To establish endometrial assembloids in vitro, pre-receptive endometrium from reproductive-age women was dissociated into single cells or small cell masses. These cells then self-assembled into assembloids induced by various small molecules, such as Noggin, EGF, FGF2, WNT-3A and R-Spondin1 in expansion medium (ExM) (Fig.1A∼B, Fig.S1A). The assembloids derived from the first generation are used for experiments (Fig.S1B). The endometrial assembloids consist of vesicle-like glands and fibrous stromal cells (Fig. S1C, Video S1∼S4). The epithelium marker E-cadherin, stromal cell marker vimentin, and endometrial gland marker FOXA2 were expressed in the cultured endometrial assembloid, which resembled the endometrium morphologically (Fig. S1D). Moreover, the endometrial assembloids exhibited substantial expression of the proliferation marker Ki67, while the apoptosis marker cleaved caspase-3 was undetectable, signifying the assembloids’ robust proliferative capacity (Fig. S1E).

Revised version: To generate endometrial assembloids in vitro, pre-receptive endometrium from reproductive-age women was dissociated into single cells or small cell masses. These cells then self-assembled into assembloids induced by various small molecules, such as Noggin, EGF, FGF2, WNT-3A and R-Spondin1 in expansion medium (ExM) (Fig.1A∼B, Fig.S1A). The assembloids derived from the first three generation are used for experiments (Fig.S1B). The endometrial assembloids comprised vesicle-like glands and fibrous stromal cells (Fig. S1C, Video S1∼S4). The epithelium marker E-cadherin, stromal cell marker vimentin, and endometrial gland marker FOXA2 were confirmed in cultured assembloid, which exhibited morphological similarity to the endometrium (Fig. S1D). Moreover, the endometrial assembloids exhibited substantial expression of the proliferation marker Ki67, with no detectable cleaved caspase-3 (apoptosis marker), indicative of strong proliferative capacity (Fig. S1E).

Developing receptive endometrial assembloids in vitro.

Added in the revised version: (A) Human endometrial assembloids constructed from adult stem cells were treated with expansion medium (ExM) (CTRL) or subjected to hormonal stimulation. Timeline of endometrial assembloid cultured by ExM (CTRL), ovarian steroid hormones simulating secretory phase (SEC), ovarian steroid hormones combining PRL and placental hormones to mimic the window of implantation (WOI). (B) Endometrial assembloids from the CTRL, SEC, and WOI groups, which were subjected to hormone treatment on Days 0, 2, and 8, exhibited comparable growth patterns throughout the culture period. Scale bar = 200 μm. (C) The dynamic changes of the counts of assembloids over time in each hormone regimen. (D) The dynamic changes of the area of assembloids over time in each hormone regimen. (E) Heatmap showing receptivity related gene expression profile of assembloids in each hormone regimen. The color represents log-transformed fold change of gene expression. (F) Validation of receptivity markers (IGFBP1, MAOA and DPP4) with immunofluorescence (IF) in the CTRL, SEC and WOI endometrial assembloids in vitro. Nuclei were counterstained with DAPI. Scale bar = 30 μm. The bar chart displaying the quantitative comparison of receptivity markers among three groups. *P ≤ 0.05, **P ≤ 0.005, ***P ≤ 0.0005, ****P ≤ 0.0001.

Removed in the revised version: (A) Human endometrial assembloids constructed from adult stem cells were treated with expansion medium (ExM) (CTRL) or subjected to hormonal stimulation. Timeline of endometrial assembloid cultured by ExM (CTRL), ovarian steroid hormones simulating secretory phase (SEC), ovarian steroid hormones combining PRL and placental hormones to mimic the window of implantation (WOI).

Removed in the revised version: (B) Endometrial assembloids in the CTRL, SEC and WOI groups displayed similar growth patterns during the culture period. Scale bar = 200 μm.

Removed in the revised version: (C) The dynamic changes of the counts of assembloids over time in each hormone regimen.

Removed in the revised version: (D) The dynamic changes of the area of assembloids over time in each hormone regimen.

Removed in the revised version: (E) Heatmap showing receptivity related gene expression profile of assembloids in each hormone regimen. The color represents log-transformed fold change of gene expression.

Removed in the revised version: (F) Validation of receptivity markers (IGFBP1, MAOA and DPP4) with immunofluorescence (IF) in the CTRL, SEC and WOI endometrial assembloids in vitro. Nuclei were counterstained with DAPI. Scale bar = 30 μm. The bar chart displaying the quantitative comparison of receptivity markers among three groups. *P ≤ 0.05, **P ≤ 0.005, ***P ≤ 0.0005, ****P ≤ 0.0001.

Earlier version: Functionally, endometrial assembloids effectively secreted glycogen into the lumen, mirroring the in vivo endometrial activity, thereby providing nourishment for embryo implantation (Fig. S1D). Moreover, after two days of estrogen (E2) treatment followed by fourteen days of medroxyprogesterone acetate (MPA) and cAMP administration, the assembloids exhibited a significant upregulation of progesterone receptor (PRA/B) expression, a modest increase in estrogen receptor α (ERα), and enhanced expression of estrogen-responsive genes EGR1 and OLFM4, along with progesterone-responsive genes PGR and PAEP at the transcriptional level (Fig. S1F∼S1G). It reflected a comparable hormonal responsiveness in the in vivo endometrium.

Revised version: Functionally, endometrial assembloids effectively secreted glycogen into the lumen, mirroring the nutrient-secreting activity of the endometrium to support embryo implantation (Fig. S1D). Moreover, after two days of estrogen (E2) treatment followed by fourteen days of medroxyprogesterone acetate (MPA) and cAMP, the assembloids exhibited significant transcriptional upregulation of progesterone receptor (PRA/B), a modest increase in estrogen receptor α (ERα), and enhanced expression of estrogen-responsive genes EGR1 and OLFM4, along with progesterone-responsive genes PGR and PAEP (Fig. S1F∼S1G). It reflected hormonal responsiveness analogous to that of the in vivo endometrium (2).

Earlier version: To identify hormonal regimens that induce implantation window, pregnancy-related hormones were supplemented into the culture system following the two-day E2. E2 and MPA promote the transition of endometrial assembloids into the secretory phase, while pregnancy hormones can promote the further differentiation. Prolactin (PRL) promotes immune regulation and angiogenesis during implantation(2, 21). Human chorionic gonadotropin (hCG) improves endometrial thickness and receptivity (22, 23). Human placental lactogen (hPL) promotes the development and function of endometrial glands(24). Hormone dosage was primarily based on peri-pregnant maternal body or localized endometrium levels (2). A comparison of multiple groups revealed similar counts, area, and average intensity of assembloids over time (Fig.1C-D, Fig.S1H). However, only the final recipe (i.e., a combination of E2, MPA, cAMP, PRL, hCG and hPL) exhibited an endometrial receptivity-related gene expression profile, which highly expressed genes positively correlated with endometrial receptivity, and lowly expressed genes negatively correlated with receptivity, compared to the other hormone formulations (Fig.1E). The assembloids induced with this scheme were defined as WOI assembloids (Fig. 1A). In contrast, endometrial assembloids maintained in ExM constituted the “control (CTRL)” group, whereas assembloids treated with E2 for two days, followed by E2, MPA, and cAMP for an additional six days, were induced to the secretory phase as previously reported(9), and designated as the “secretory (SEC)” group (Fig. 1A). There was no significant difference in the morphology of assembloids among the three groups (Fig.1B), but when comparing receptivity markers, WOI assembloids exhibited high expression of IGFBP1, MAOA, and DPP4 (Fig. 1F), with increased glycogen secretion (Fig.S1I). Theoretically, the WOI assembloids are initially in the secretory phase, thus sharing characteristics with the SEC assembloids. More importantly, they are expected to display traits typical of the mid-secretory phase.

Revised version: To identify hormone regimens that induce implantation window, pregnancy-related hormones were supplemented into the culture system following 2 days of E2 priming. E2 and MPA drive the transition of endometrial assembloids to the secretory phase, while pregnancy hormones promote further differentiation. Prolactin (PRL) promotes immune regulation and angiogenesis during implantation(2, 21). Human chorionic gonadotropin (hCG) improves endometrial thickness and receptivity (22, 23). Human placental lactogen (hPL) promotes the development and function of endometrial glands(24). Hormone dosages were primarily based on peri-pregnant maternal systemic or local endometrium levels (2). Multigroup comparison revealed similar counts, area, and average intensity of assembloids over time (Fig.1C-D, Fig.S1H). However, only the final cocktail (i.e., a combination of E2, MPA, cAMP, PRL, hCG and hPL) exhibited an endometrial receptivity-related gene expression profile, with high expression of receptivity-promoting genes and low expression of receptivity-inhibiting genes relative to other hormone formulations (Fig.1E). The assembloids generated with this regimen were defined as WOI assembloids (Fig. 1A). For controls, assembloids maintained in ExM served as the “control (CTRL)” group, while those treated with 2 days of E₂ followed by 6 days of E₂, MPA, and cAMP were induced to the secretory phase (as previously described (9)) and designated the “secretory (SEC)” group (Fig. 1A). There was no significant difference in the morphology of assembloids among the three groups (Fig.1B), but WOI assembloids exhibited elevated expression of the receptivity markers IGFBP1, MAOA, and DPP4 (Fig.1F), and increased glycogen secretion (Fig.S1I). Theoretically, the WOI assembloids originate from the secretory phase and thus share characteristics with SEC assembloids, but crucially, they are expected to exhibit hallmark features of the mid-secretory phase.

Receptive endometrial assembloids mimicked the implantation - window endometrium

Earlier version: Single-cell transcriptomics analysis identified the presence of epithelium, stromal cells, and immune cells in WOI assembloids with reference to CellMarker, PanglaoDB, Human Cell Atlas, Human Cell Landscape, and scRNASeqDB, and previous endometrium related studies(1, 9, 11, 25) (Fig. 2A, Fig. S2A∼S2B). The WOI assembloids exhibited similarities to the mid-secretory endometrium in terms of glandular and luminal epithelium, secretory epithelium, LGR5 epithelium, ciliated epithelium, and EMT-derived cells, as evidenced by comparisons of scRNA-seq data from our assembloids and the mid-secretory endometrium as described by Stephen R. Quake in 2020 (1) (Fig.2A, Fig.S2C∼S2D). The morphology of immune and stromal cells was analyzed through 3D clearing staining and light sheet microscopy imaging, with vimentin labeling stromal cells (Vimentin+ or Vimentin+ F-actin+), CD45 and CD44 indicating immune cells, and FOXA2 identifying glands (Fig.2B∼2D). Furtherly, we confirmed the presence and types of immune cells using flow cytometry. White blood cells (WBC) were identified as CD45+ cells, with T cells, macrophages and NK cells characterized as CD45+CD3+ cells, CD45+CD68+CD11b+ cells and CD56+CD16- cells, respectively (Fig. 2E).

Revised version: Single-cell transcriptomics analysis, with reference to CellMarker, PanglaoDB, Human Cell Atlas, Human Cell Landscape, and scRNASeqDB, and prior endometrium related studies(1, 9, 11, 25), identified the presence of epithelial, stromal, and immune cells within WOI assembloids (Fig. 2A, Fig. S2A∼S2B). Comparative analysis of scRNA-seq data from our assembloids and mid-secretory endometrial tissue (as described by Stephen R. Quake et al. in 2020 (1) and Garcia-Alonso in 2021 (25)) revealed that WOI assembloids exhibited similarities to the mid-secretory endometrium in terms of glandular epithelium, ciliated epithelium, and stromal cells (Fig.2A, Fig.S2C∼S2F). The morphology of immune and stromal cells was visualized via 3D clearing staining and light sheet microscopy imaging, with vimentin labeling stromal cells (Vimentin+ or Vimentin+ F-actin+), CD45 and CD44 indicating immune cells, and FOXA2 identifying glands (Fig.2B∼2D). Furthermore, flow cytometry was employed to validate the presence and subset composition of immune cells. White blood cells (WBC) were identified as CD45+ cells, with T cells, macrophages and NK cells characterized as CD45+CD3+, CD45+CD68+CD11b+ and CD56+CD16- cells, respectively (Fig. 2E). WOI assembloids displayed characteristic features of the receptive endometrium. On the one hand, they secreted significantly higher levels of glycogen into the lumen compared to other groups (Fig. S1I). On the other hand, they possessed various characteristic microstructure of implantation window, including elongated microvilli and increased glycogen, pinopodes, and especially cilia (Fig.2F).

Receptive endometrial assembloids mimicked the implantation - window endometrium.

Added in the revised version: (A) T-SNE plot of scRNA-seq data from three individual endometrial assembloids of the CTRL, SEC and WOI groups (left). T-SNE plot of combined scRNA-seq data from the three kinds of assembloids and mid-secretory endometrium (right). (B) Exhibition of stromal cell marked by vimentin of CTRL assembloid through whole-mount clearing, immunostaining and light sheet microscopy imaging. Nuclei were counterstained with DAPI. The arrowhead indicates stromal cells. Scale bar = 40 μm (left), Scale bar = 30 μm (right). (C) Whole-mount immunofluorescence showed that Vimentin+ F-actin+ cells (stromal cells) were arranged around the glandular spheres that were only F-actin+. Scale bar = 50 μm. (D) Exhibition of immune cell marked by CD45 and CD44, and endometrial gland marked by FOXA2 of CTRL assembloid through whole-mount clearing, immunostaining and light sheet microscopy imaging. Nuclei were counterstained with DAPI. The arrowhead indicates immune cells. Scale bar = 50 μm (left), Scale bar = 10 μm (right). (E) Flow cytometric analysis of T cells and macrophages in the CTRL endometrial assembloid. Gating strategy used for determining white blood cells (WBC) (CD45+ cells), T cells (CD45+CD3+ cells) and macrophages (CD45+CD68+CD11b+ cells). (F) Electron micrograph of the CTRL (top), SEC (middle) and WOI (bottom) endometrial assembloid showing pinopodes (P), glycogen granule (asterisk), microvilli (white arrows) and cilia (orange arrows). Scale bar = 1 μm. Quantitative comparison of pinopodes, glycogen, microvilli, and cilia in the CTRL, SEC and WOI assembloids. *P ≤ 0.05, **P ≤ 0.005, ***P ≤ 0.0005, ****P ≤ 0.0001. (G) Heatmap and bubble diagram illustrating highly expressed genes as well as GO functions enriched in both assembloids during the WOI and mid secretory endometrial tissue in terms of SOX9+ proliferative epithelium, stem-derived epithelium, secretory epithelium, proliferative epithelium, unciliated epithelium, stromal cells and EMT-derived cells. The color of heatmap represents log-transformed fold change of gene expression. (H) Heatmaps showing differentially expressed TFs of endometrial assembloids and endometrium in the secretory epithelium (left) and EMT-derived cells (right). The color represents log-transformed fold change of gene expression.

Removed in the revised version: (A) T-SNE plot of scRNA-seq data from three individual endometrial assembloids of the CTRL, SEC and WOI groups (left). T-SNE plot of combined scRNA-seq data from the three kinds of assembloids and mid-secretory endometrium (right).

Removed in the revised version: (B) Exhibition of stromal cell marked by vimentin of CTRL assembloid through whole-mount clearing, immunostaining and light sheet microscopy imaging. Nuclei were counterstained with DAPI. The arrowhead indicates stromal cells. Scale bar = 40 μm (left), Scale bar = 30 μm (right).

Removed in the revised version: (C) Whole-mount immunofluorescence showed that Vimentin+ F-actin+ cells (stromal cells) were arranged around the glandular spheres that were only F-actin+. Scale bar = 50 μm.

Removed in the revised version: (D) Exhibition of immune cell marked by CD45 and CD44, and endometrial gland marked by FOXA2 of CTRL assembloid through whole-mount clearing, immunostaining and light sheet microscopy imaging. Nuclei were counterstained with DAPI. The arrowhead indicates immune cells. Scale bar = 50 μm (left), Scale bar = 10 μm (right).

Removed in the revised version: (E) Flow cytometric analysis of T cells and macrophages in the CTRL endometrial assembloid. Gating strategy used for determining white blood cells (WBC) (CD45+ cells), T cells (CD45+CD3+ cells) and macrophages (CD45+CD68+CD11b+ cells).

Removed in the revised version: (F) Electron micrograph of the CTRL (top), SEC (middle) and WOI (bottom) endometrial assembloid showing pinopodes (P), glycogen granule (asterisk), microvilli (white arrows) and cilia (orange arrows). Scale bar = 1 μm. Quantitative comparison of pinopodes, glycogen, microvilli, and cilia in the CTRL, SEC and WOI assembloids. *P ≤ 0.05, **P ≤ 0.005, ***P ≤ 0.0005, ****P ≤ 0.0001.

Removed in the revised version: (G) Heatmap and bubble diagram illustrating highly expressed genes as well as GO functions enriched in both assembloids during the WOI and mid secretory endometrial tissue in terms of SOX9+ proliferative epithelium, stem-derived epithelium, secretory epithelium, proliferative epithelium, unciliated epithelium, stromal cells and EMT-derived cells. The color of heatmap represents log-transformed fold change of gene expression.

Removed in the revised version: (H) Heatmaps showing differentially expressed TFs of endometrial assembloids and endometrium in the secretory epithelium (left) and EMT-derived cells (right). The color represents log-transformed fold change of gene expression.

Removed in the revised version: The WOI assembloids displayed characteristic features of the receptive endometrium. On the one hand, the WOI assembloids secreted more glycogen into the lumen (Fig. S1I). On the other hand, the WOI assembloids possessed various characteristic microstructure of implantation window, including elongated microvilli and increased glycogen, pinopodes, and especially cilia (Fig.2F).

Earlier version: Further, we concentrated on the comparison of gene expression and transcriptional regulation during the mid-secretory phase, as this phase marks the implantation window and involves substantial transcriptional alterations (Fig.2G∼2H). We referred to the scRNA-seq data of the mid-secretory phase from Stephen R. Quake 2020 (1). Pathways related to mitochondrial energy metabolism and cell adhesion exhibited upregulation in both the WOI assembloid and mid-secretory endometrium in comparison to the CTRL and SEC assembloids (Fig. 2G). The crucial transcription factors (TFs) found in the secretory epithelium and EMT-derived cells, which are implicated in implantation, revealed similarities between WOI assembloids and mid-secretory endometrium. Specifically, the secretory epithelium exhibited comparable TFs related to hypoxia response (such as previously reported HIF1A(26)), embryo implantation (such as FBLN1(27)), lipid metabolism (such as VMP1(28)), cell migration, and cell junction (such as TJP1(29)) (Fig. 2H). Similarly, EMT-derived cells also expressed similar TFs involved in endometrial decidualization (such as S100A10(30)), EMT (such as FAT1(31) and FOXF2(32)), and receptivity (such as NEAT1(33), SERPINB9(34), SOX17(35) and SOX4(36)) (Fig. 2H).

Revised version: We next focused on transcriptional profiling and regulation during the mid-secretory phase, as this phase marks the implantation window and involves substantial transcriptional alterations (Fig.2G∼2H). For this analysis, we referred scRNA-seq data of the mid-secretory phase from Stephen R. Quake 2020 (1). Pathways related to mitochondrial energy metabolism and cell adhesion were significantly upregulated in both WOI assembloid and mid-secretory endometrium, relative to CTRL and SEC assembloids (Fig. 2G). Furthermore, key transcription factors (TFs) associated with implantation—identified in secretory epithelial cells and EMT-derived cells—exhibited highly conserved expression patterns between WOI assembloids and the mid-secretory endometrium. Specifically, secretory epithelium shared comparable TFs linked to hypoxia response (e.g., HIF1A(26)), embryo implantation (e.g., FBLN1(27)), lipid metabolism (e.g., VMP1(28)), cell migration, and cell junction (e.g., TJP1(29)) (Fig. 2H). Similarly, EMT-derived cells also expressed conserved TFs involved in endometrial decidualization (e.g., S100A10(30)), EMT (e.g., FAT1(31) and FOXF2(32)), and receptivity (e.g., NEAT1(33), SERPINB9(34), SOX17(35) and SOX4(36)) (Fig. 2H).

Earlier version: We then conducted the endometrial receptivity test (ERT) (37) to further assess the receptive state of WOI assembloids. ERT is a kind of gene analysis-based method for detecting endometrial receptivity, which combines high-throughput sequencing and machine learning to analyze the expression of endometrial receptivity-related genes (37). It is currently used in clinical practice to determine endometrial receptivity and guide personalized embryo transfer. The WOI assembloids derived from pre-receptive endometrium were observed to transit into the receptivity phase (Fig. S1J).

Revised version: To further validate the receptive phenotype of WOI assembloids, we performed Endometrial Receptivity Analysis (ERA), a gene expression profiling-based diagnostic assay that integrates high-throughput sequencing and machine learning to quantify the expression of endometrial receptivity-associated genes (37). Currently employed in clinical practice to assess endometrial receptivity and guide personalized embryo transfer, ERA revealed that WOI assembloids derived from pre-receptive endometrial tissue successfully transitioned to a receptive state (Fig. S1J).

Earlier version: Collectively, these results demonstrated that the WOI assembloids closely resemble the in vivo endometrium during the implantation window in terms of structures and molecular characteristics.

Revised version: Collectively, these findings confirm that WOI assembloids closely recapitulate the structural and molecular characteristics of the in vivo endometrium during the implantation window.

Receptive endometrial assembloids recapitulate WOI-associated hormone response

Earlier version: We analyzed the transcriptome and proteome profiles of WOI assembloids compared to those of CTRL or SEC assembloids to assess WOI associated biological characteristics (Fig. 3A, Fig. S3A∼S3E).

Revised version: To characterize the biological hallmarks associated with the implantation window, we performed integrated transcriptomic and proteomic profiling of WOI assembloids, with CTRL and SEC assembloids (Fig. 3A, Fig. S3A∼S3E).

Receptive endometrial assembloids recapitulate WOI-associated hormone response.

Added in the revised version: (A) Principal component analysis (PCA) plot computed with differentially expressed genes in the bulk transcriptome of endometrial assembloids belonging to the CTRL, SEC and WOI groups. (B) Heatmap showing that enrichment of differentially expressed genes for the terms of hormone response, monocarboxylic acid metabolism, lipid metabolism, and negative regulation of cell differentiation. The color represents log-transformed fold change of gene expression. (C) Responsiveness to progesterone and estrogen was evaluated by IF to PRA/B and OLFM4 with IF respectively. Scale bar = 40 μm, **P ≤ 0.005. (D) Exhibition of implantation marker (FOXO1) and endometrial gland marker (FOXA2) through combination of assembloid clearing, IF and light sheet microscopy. Nuclei were counterstained with DAPI. **P ≤ 0.005. (E) Bar graph exhibiting various percentages of each cell type in the three groups. (F) GSEA between the SEC and WOI groups for secretory epithelium. (G) Pseudotime trajectory showing the transformation between proliferative and secretory epithelium in the CTRL, SEC and WOI groups. Arrows indicate the direction of the pseudotime trajectory. The black dot indicates the key differentiation node.

Removed in the revised version: (A) Principal component analysis (PCA) plot computed with differentially expressed genes in the bulk transcriptome of endometrial assembloids belonging to the CTRL, SEC and WOI groups.

Removed in the revised version: (B) Heatmap showing that enrichment of differentially expressed genes for the terms of hormone response, monocarboxylic acid metabolism, lipid metabolism, and negative regulation of cell differentiation. The color represents log-transformed fold change of gene expression.

Removed in the revised version: (C) Responsiveness to progesterone and estrogen was evaluated by IF to PRA/B and OLFM4 with IF respectively. Scale bar = 40 μm, **P ≤ 0.005.

Removed in the revised version: (D) Exhibition of implantation marker (FOXO1) and endometrial gland marker (FOXA2) through combination of assembloid clearing, IF and light sheet microscopy. Nuclei were counterstained with DAPI. **P ≤ 0.005.

Removed in the revised version: (E) Bar graph exhibiting various percentages of each cell type in the three groups.

Removed in the revised version: (F) GSEA between the SEC and WOI groups for secretory epithelium.

Removed in the revised version: (G) Pseudotime trajectory showing the transformation between proliferative and secretory epithelium in the CTRL, SEC and WOI groups. Arrows indicate the direction of the pseudotime trajectory. The black dot indicates the key differentiation node.

Earlier version: WOI assembloids exhibited a robust hormone response, as demonstrated by the upregulation of PGR at the transcriptome level (Fig. 3B). Integrated analysis of the transcriptome and proteome data identified 179 upregulated genes/proteins in WOI assembloids compared to the CTRL group, with most of these genes/proteins implicated in the estrogen response (Fig. S3F∼S3G). Furthermore, the analysis of progesterone response levels using PRA/B immunostaining indicated the highest levels in the WOI group, accompanied by an upregulation of the estrogen responsive protein OLFM4 (Fig. 3C). FOXO1, a crucial marker of endometrial receptivity reliant on PGR signaling, exhibited significantly elevated expression in WOI assembloids compared to CTRL assembloids (Fig.3D, Video S5), suggesting the involvement of progesterone signaling in the establishment of WOI assembloids. These findings collectively demonstrated a strong response of WOI assembloids to estrogen and progesterone.

Revised version: WOI assembloids exhibited a robust hormone response, as evidenced by upregulated PGR expression at the transcriptome level (Fig. 3B). Integrated multi-omics analysis identified 179 genes/proteins that were significantly upregulated in WOI versus CTRL assembloids, the majority of which are implicated in estrogen signaling pathways (Fig. S3F∼S3G). Furthermore, immunostaining for PRA/B, which was used to quantify progesterone responsiveness, revealed the highest signal intensity in the WOI group, concurrent with upregulation of the estrogen-responsive protein OLFM4 (Fig. 3C). FOXO1, a crucial marker of endometrial receptivity reliant on PGR signaling, was significantly elevated in WOI compared to CTRL assembloids (Fig.3D, Video S5). These observations suggest that progesterone signaling is central to the establishment of the WOI phenotype in assembloids, and collectively demonstrate that WOI assembloids mount a robust response to both estrogen and progesterone.

Earlier version: Hormone response increased secretory epithelium and decreased proliferative epithelium in the SEC and WOI assembloids, suggesting the transformation from proliferation to secretory phase (Fig. 3E). The secretory epithelium, critical for the implantation window, contributed to cellular metabolic processes and HIF-1 signaling pathway response to hypoxia at the single-cell level (Fig. S2E∼S2F). The secretory epithelium of WOI assembloids exhibited enhanced peptide metabolism and mitochondrial energy metabolism compared to the SEC group, supporting endometrial decidualization and embryo implantation (Fig. 3F). Proliferative epithelium differentiates into secretory epithelium under the regulation of branch nodal genes between state 5 and state 6, such as the KRT19, MALAT1, MT2A, RPL and RPS families, as revealed by single-cell trajectory (Fig. 3G, Fig. S2H∼S2J). The WOI assembloids showed more thoroughly differentiation from proliferative to secretory epithelium in relation to the SEC assembloids (Fig. 3G).

Revised version: Hormonal stimulation induced an expansion of secretory epithelium and a concomitant reduction in proliferative epithelium in SEC and WOI assembloids, indicative of the proliferative-to-secretory phase transition (Fig. 3E). At the single-cell level, the secretory epithelium, which is critical for implantation window, was enriched for genes involved in cellular metabolic processes and hypoxic responses (HIF-1 signaling pathway) (Fig. S2G∼S2H). Compared to SEC assembloids, secretory epithelium in WOI assembloids exhibited enhanced peptide metabolism and mitochondrial energy metabolism (Fig. 3F), which are functional adaptations that support endometrial decidualization and embryo implantation. Single-cell trajectory analysis further revealed that proliferative epithelium differentiates into secretory epithelium under the regulation of key nodal genes defining the transition between states 5 and 6, including KRT19, MALAT1, MT2A, and the RPL and RPS families (Fig. 3G, Fig. S2J∼S2L). Notably, WOI assembloids displayed a more complete proliferative-to-secretory epithelial transition than SEC assembloids (Fig. 3G).

Earlier version: Overall, the WOI assembloids possessed the hormone response characteristic of implantation window, closely resembling the gene traits of embryo implantation.

Revised version: Overall, WOI assembloids possessed the hormone response characteristic of implantation window, closely resembling the gene traits of embryo implantation.

Receptive endometrial assembloids possess enhanced energy metabolism

Earlier version: The WOI assembloids exhibited upregulation of monocarboxylic acid and lipid metabolism (represented by SLC25A1(38)), and hypoxia response (represented by HIF1α(26)) (Fig.3B, Fig.S3G∼S3I). Likewise, the secretory epithelium, critical for the implantation window, contributed to cellular metabolic processes and HIF-1 signaling pathway response to hypoxia at the single-cell level (Fig. S2E∼S2F).

Revised version: WOI assembloids exhibited upregulation of monocarboxylic acid and lipid metabolism (represented by SLC25A1(38)), and hypoxia response (represented by HIF1α(26)) (Fig.3B, Fig.S3G∼S3I). Likewise, the secretory epithelium, critical for the implantation window, contributed to cellular metabolic processes and HIF-1 signaling pathway response to hypoxia at the single-cell level (Fig. S2G∼S2H).

Earlier version: To further investigate this trend, the Mfuzz algorithm was utilized to analyze gene expression across these three groups, focusing on gene clusters that were progressively upregulated or downregulated. It was observed that mitochondrial genes exhibited the highest expression levels in WOI endometrial assembloids (Fig. 4A). Additionally, at the protein level, WOI endometrial assembloids showed sustained high expression of mitochondrial proteins compared to SEC assembloids (Fig. 4B). TEM analysis revealed that WOI endometrial assembloids had the largest average mitochondrial area (Fig. 4C). The expression of mitochondrial-related genes increased from CTRL to SEC to WOI assembloids, with COA1 ensuring proper nuclear-mitochondrial connection(39), OXA1L promoting mitochondrial translation(40), and TIMMDC1 being crucial for mitochondrial complex I assembly(41) (Fig. 4D). The WOI assembloids notably expressed higher levels of OXA1L and TIMMDC1 than the SEC assembloids (Fig. 4D). Furthermore, WOI assembloids produced more ATP and IL8(42) (Fig. 4E).

Revised version: To further investigate this trend, the Mfuzz algorithm was utilized to analyze gene expression across these three groups, focusing on gene clusters that were progressively upregulated or downregulated. Mitochondrial-related genes were found to exhibit the highest expression levels in WOI endometrial assembloids (Fig. 4A). Concordantly, at the protein level, WOI endometrial assembloids maintained significantly higher expression of mitochondrial proteins compared to SEC assembloids (Fig. 4B). TEM analysis further revealed that WOI endometrial assembloids possessed the largest average mitochondrial area among the three groups (Fig. 4C). The expression of mitochondrial-related genes increased from CTRL to SEC to WOI assembloids, with COA1 ensuring proper nuclear-mitochondrial connection(39), OXA1L promoting mitochondrial translation(40), and TIMMDC1 being crucial for mitochondrial complex I assembly(41) (Fig. 4D). WOI assembloids notably expressed higher levels of OXA1L and TIMMDC1 than the SEC assembloids (Fig. 4D). Furthermore, WOI assembloids produced more ATP and IL8(42) (Fig. 4E).

Receptive endometrial assembloids possess enhanced energy metabolism.

Earlier version: (A) The Mfuzz trend analysis displayed the transcriptional variation trends of five clusters from CTRL, SEC, and WOI groups (with a focus on the differences between SEC and WOI assembloids). The heatmap illustrated corresponding gene expression profile (where color represents Z-score). The bubble plot showed the associated GO functions (with bubble size representing the number of genes and bubble color indicating the P value).

Revised version: (A) The Mfuzz trend analysis displayed the transcriptional variation trends of five clusters from CTRL, SEC, and WOI groups (with a focus on the differences between SEC and WOI assembloids). The heatmap illustrated corresponding gene expression profile (where color represents Z-score). The bubble plot showed the associated GO functions (with bubble size representing the number of genes and bubble color indicating the P value). (B) The circular heatmap illustrated the functional differences of SEC and WOI assembloids at the proteomic level. The color represents protein expression levels, and the innermost circle color represents GO functions. (C) Transmission electron microscopy images displayed the mitochondrial morphology of CTRL, SEC, and WOI assembloids, along with a quantitative comparison of mitochondrial area. Scale bar = 1 μm. (D) RT-qPCR assessed the expression levels of mitochondrial function-related genes in the three assembloid groups. (E) Quantitative comparison of the concentration of ATP (left) and IL8 (right) released by CTRL, SEC and WOI assembloids. *P<0.05,**P<0.005,***P<0.0005,****P<0.0001.

Removed in the revised version: (B) The circular heatmap illustrated the functional differences of SEC and WOI assembloids at the proteomic level. The color represents protein expression levels, and the innermost circle color represents GO functions.

Removed in the revised version: (C) Transmission electron microscopy images displayed the mitochondrial morphology of CTRL, SEC, and WOI assembloids, along with a quantitative comparison of mitochondrial area. Scale bar = 1 μm.

Removed in the revised version: (D) RT-qPCR assessed the expression levels of mitochondrial function-related genes in the three assembloid groups.

Removed in the revised version: (E) Quantitative comparison of the concentration of ATP (left) and IL8 (right) released by CTRL, SEC and WOI assembloids. *P<0.05<**P<0.005<***P<0.0005<****P<0.0001.

Earlier version: Thus, compared to SEC assembloids, WOI assembloids exhibited increased energy metabolism with larger and enhanced functional mitochondria.

Revised version: Thus, compared to SEC assembloids, WOI assembloids exhibited enhanced energy metabolism, underpinned by larger, functionally competent mitochondria.

Receptive endometrial assembloids increased the ciliary assembly and motility

Earlier version: The growth and development, assembly and movement of cilia, a characteristic endometrial structure, are essential for the establishment of endometrial implantation window and embryo implantation. Under the electron microscopy, the cilia were most observed in the WOI assembloids (Fig. 2E). Cilia-related genes and proteins exhibited the highest expression levels in WOI endometrial assembloids (Fig. 5A∼5B). Compared to SEC assembloids, WOI endometrial assembloids showed upregulation in ciliary assembly, ciliary basal body, and motile cilia-related gene expressions, while genes related to non-motile cilia were downregulated (Fig. 5A).

Revised version: Cilia are specialized structural components of the endometrium, whose growth and development, assembly and movement are essential for establishing the endometrial implantation window and facilitating embryo implantation. Under the electron microscopy, cilia were most observed in the WOI assembloids (Fig. 2E). Concordantly, cilia-related genes and proteins exhibited peak expression levels in WOI endometrial assembloids (Fig. 5A∼5B). Transcriptomic profiling further demonstrated that, relative to SEC assembloids, WOI assembloids displayed upregulated expression of genes linked to ciliary assembly, ciliary basal body, and motile cilia, whereas genes associated with non-motile cilia were downregulated (Fig. 5A). Key differentially expressed genes in WOI assembloids included NEK2 (ciliary assembly), CFAP36 (ciliary basal body), and DNAH9 and TPPP (motile cilia), with NEK2 showing the most pronounced upregulation (Fig.5B). At the protein level, WOI assembloids also exhibited elevated expression of factors involved in ciliary assembly (TBC1D1, IFT22, IFT57), ciliary basal body biogenesis (PJA2), and motile cilia function (Fig. 5C).

Receptive endometrial assembloids increased the ciliary assembly and motility.

Added in the revised version: (A) The heatmap illustrated the expression of cilia-related genes across the CTRL, SEC, and WOI assembloids. The color represents Z-score, while the leftmost block indicates various characteristic functions related to cilia. (B) RT-qPCR assessed the expression levels of cilia-related genes in the three assembloids. (C) The histogram showed the expression of cilia-related proteins in the three groups of assembloids. The color of the longitudinal protein names corresponds to the color of cilia-related functional blocks in Fig 5A. (D) IF analysis of cilia assembly marked by acetyl-α-tubulin. Nuclei were counterstained with DAPI. The arrowhead indicates cilia. Scale bar = 50 μm. (E) GSEA between the SEC and WOI groups for ciliated epithelium. (F) Pseudotime trajectory showing the transformation between ciliated and unciliated epithelium in the CTRL, SEC and WOI groups. Arrows indicate the direction of the pseudotime trajectory. (G) Dot plots demonstrating the Cellphone DB analysis of relevant receptors and ligands of ciliated epithelium with other cell types. The size of the dot represents the level of significance. The color of the dot indicates the mean of the average expression level of interacting molecule 1 in ciliated epithelium and molecule 2 in other cell types. (H) Proximity ligation assay (PLA) validating the interactions of ROR2-Wnt5A and CD74-COPA in the CTRL, SEC and WOI assembloids. Red signals the interaction of two proteins. Nuclei were counterstained with DAPI. Scale bar = 20 μm. *P<0.05,**P<0.005,***P<0.0005,****P<0.0001.

Added in the revised version: Additionally, acetyl-α-tubulin (cilia marker (15)) were highly expressed in the WOI assembloids (Fig. 5D).

Removed in the revised version: (A) The heatmap illustrated the expression of cilia-related genes across the CTRL, SEC, and WOI assembloids. The color represents Z-score, while the leftmost block indicates various characteristic functions related to cilia.

Removed in the revised version: (B) RT-qPCR assessed the expression levels of cilia-related genes in the three assembloids.

Removed in the revised version: (C) The histogram showed the expression of cilia-related proteins in the three groups of assembloids. The color of the longitudinal protein names corresponds to the color of cilia-related functional blocks in Fig 5A.

Removed in the revised version: (D) IF analysis of cilia assembly marked by acetyl-α-tubulin. Nuclei were counterstained with DAPI. The arrowhead indicates cilia. Scale bar = 50 μm.

Removed in the revised version: (E) GSEA between the SEC and WOI groups for ciliated epithelium.

Removed in the revised version: (F) Pseudotime trajectory showing the transformation between ciliated and unciliated epithelium in the CTRL, SEC and WOI groups. Arrows indicate the direction of the pseudotime trajectory.

Removed in the revised version: (G) Dot plots demonstrating the Cellphone DB analysis of relevant receptors and ligands of ciliated epithelium with other cell types. The size of the dot represents the level of significance. The color of the dot indicates the mean of the average expression level of interacting molecule 1 in ciliated epithelium and molecule 2 in other cell types.

Removed in the revised version: (H) Proximity ligation assay (PLA) validating the interactions of ROR2-Wnt5A and CD74-COPA in the CTRL, SEC and WOI assembloids. Red signals the interaction of two proteins. Nuclei were counterstained with DAPI. Scale bar = 20 μm. *P<0.05<**P<0.005<***P<0.0005<****P<0.0001.

Removed in the revised version: Transcriptomic analysis further confirmed the upregulation of genes involved in ciliary assembly (NEK2), ciliary basal body (CFAP36), and motile cilia (DNAH9 and TPPP) in WOI endometrial assembloids, with NEK2 showing the most significant difference (Fig.5B). Meanwhile, expressions of proteins related to ciliary assembly, ciliary basal body, and motile cilia were also increased in WOI assembloids, with the most notable differences observed in TBC1D1, IFT22, and IFT57 (ciliary assembly), and PJA2 (ciliary basal body) (Fig. 5C). Acetyl-α-tubulin (cilia marker (15)) were highly expressed in the WOI assembloids (Fig. 5D).

Earlier version: Single-cell transcriptome analysis revealed hormone treatment increased ciliated epithelium and decreased unciliated epithelium in SEC and WOI groups (Fig. 3E). Ciliated epithelium functioned in protein binding, cilium organization and assembly, while unciliated epithelium acted on actin cytoskeleton and translation (Fig. S2E∼F). The WOI assembloids’ ciliated epithelium regulated vasculature development and displayed higher transcriptional activity than the SEC group (Fig. 5E). Ciliated-unciliated epithelium transition occurred during the menstrual cycle, which was regulated by key genes, such as GAS5, JUN, RPL and RPS families (Fig. 5F, Fig.S2K∼S2M). The implementation of aforementioned functions depended on its interaction with other cells. CellPhoneDB is a useful tool to investigate ligand-receptor interactions between the cells(43). Ciliated epithelium of WOI assembloids interacted with immune cells and secretory epithelium, showing enhanced invasion ability via CD74-COPA(44), and ROR2-WNT5A(45), which was validated by proximity ligation assay (PLA) (Fig. 5G∼5H).

Revised version: Single-cell transcriptome analysis revealed hormone treatment increased ciliated epithelium and decreased unciliated epithelium in SEC and WOI groups (Fig. 3E). Ciliated epithelium functioned in protein binding, cilium organization and assembly, while unciliated epithelium acted on actin cytoskeleton and translation (Fig. S2E∼F). Notably, WOI assembloids’ ciliated epithelium additionally regulated vasculature development and displayed higher transcriptional activity than the SEC group (Fig. 5E). The ciliated to unciliated epithelial transition occurred during the menstrual cycle, which was regulated by key genes, such as GAS5, JUN, RPL and RPS families (Fig. 5F, Fig.S2M∼S2O). Given that these functional specializations are likely mediated by intercellular crosstalk, we employed CellPhoneDB (a computational tool for predicting ligand–receptor interactions (43)) to investigate cell–cell communication networks. Ciliated epithelium of WOI assembloids interacted with immune cells and secretory epithelium, showing enhanced invasion ability via CD74-COPA(44), and ROR2-WNT5A(45), which was validated by proximity ligation assay (PLA) (Fig. 5G∼5H).

In summary, the WOI assembloids revealed accumulated ciliated epithelium’s role in preparing the implantation window, with increased ciliary assembly and motility compared to SEC assembloids.

Receptive endometrial assembloids experienced epithelial-mesenchymal transition (EMT)

Earlier version: The WOI assembloids displayed upregulated cell differentiation not only at assembloid level (Fig. 3B) but also at cellular level, which is represented by EMT. EMT is a common and crucial biological event in the endometrium during the implantation window(46). During the EMT process, epithelial cells lose their epithelial characteristics while gaining migratory and invasive properties of fibroblasts. Synthetic analysis of the transcriptome and proteome revealed increased EMT in WOI assembloids (Fig. S3G).

Revised version: The WOI assembloids displayed upregulated cell differentiation not only at assembloid level (Fig. 3B) but also at cellular level, which is represented by EMT. EMT is a common and crucial biological event in the endometrium during the implantation window(46). During the EMT process, epithelial cells lose their epithelial characteristics while gaining migratory and invasive properties of fibroblasts. Integrated transcriptomic and proteomic analyses revealed increased EMT in WOI assembloids (Fig. S3G).

Earlier version: EMT-derived cells, exhibiting gene expression patterns typical of epithelial and stromal cells, as well as EMT, are more abundant in the SEC and WOI groups, and act on protein binding, cell cycle, organelle organization, and reproduction (Fig. S2E). They performed enhanced lamellipodium-mediated cell migration, cell junction and cytoskeleton regulation in the WOI assembloids compared to the SEC assembloids (Fig. S4A). The WOI assembloids exhibited more thoroughly differentiation from proliferative epithelium to EMT-derived cells than the SEC assembloids, which was regulated by key genes such as DOC2B, FXYD3, and LPCAT3 (Fig. S4B∼S4D). EMT-derived cells and epithelium cooperated during the implantation window (Fig. S4E). We found that NRP1 and SLC7A1 were highly expressed by EMT-derived cells, and their receptors (SEMA3A and CSF1) were more upregulated in the epithelium of WOI group than SEC group (Fig. S4E, S4G). NRP1-SEMA3A has been reported to promote vascularization and responds to hypoxia(47). SLC7A1(48) and CSF1(49) both support receptivity establishment, embryo implantation and development. Compared with common stromal cells, EMT-derived cells communicate slightly differently with epithelial or immune cells. CD44 and CD46, known for their roles in cell adhesion (51) and immunoregulation(52), are highly expressed in stromal cells and EMT-derived cells, respectively, and bind separately with SPP1 and JAG1 in epithelial and stromal cells (Fig. S4E∼S4G).

Revised version: EMT-derived cells, exhibiting gene expression patterns typical of epithelial and stromal cells, as well as EMT, are more abundant in the SEC and WOI groups, and act on protein binding, cell cycle, organelle organization, and reproduction (Fig. S2E). They performed enhanced lamellipodium-mediated cell migration, cell junction and cytoskeleton regulation in the WOI assembloids compared to the SEC assembloids (Fig. S4A). Single-cell trajectory analysis further revealed that WOI assembloids underwent a more complete differentiation from proliferative epithelium to EMT-derived cells than SEC assembloids, a process governed by key regulatory genes including DOC2B, FXYD3, and LPCAT3 (Fig. S4B∼S4D). EMT-derived cells and epithelium coordinate functionally during the implantation window (Fig. S4E). Specifically, EMT-derived cells highly expressed NRP1 and SLC7A1, while their receptors (SEMA3A and CSF1) were more upregulated in the epithelium of WOI versus SEC assembloids (Fig. S4E, S4G). NRP1-SEMA3A has been reported to promote vascularization and responds to hypoxia(47). SLC7A1(48) and CSF1(49) both support receptivity establishment, embryo implantation and development. Compared with canonical stromal cells, EMT-derived cells exhibited distinct patterns of crosstalk with epithelial or immune cells. For instance, CD44 and CD46, known for their roles in cell adhesion (50) and immunoregulation(51), are highly expressed in stromal cells and EMT-derived cells, respectively, and bind separately with SPP1(in epithelial cells) and JAG1 (in stromal cells) (Fig. S4E∼S4G).

Added in the revised version: In general, WOI assembloids recapitulate the EMT-driven endometrial remodeling process that underpins the transition to the implantation window.

Removed in the revised version: In general, WOI assembloids demonstrated EMT’s role in mediating endometrial transformation towards the implantation window.

The receptive endometrial assembloids possess the potential for embryo implantation

Earlier version: We further investigated the key biological function of the implantation window, specifically embryo implantation, using the assembloids. Given the rarity and ethical constraints associated with human embryos, we employed blastoids (corresponding to the human embryo at 6 days post-fertilization, referred to as “Day 6”) for implantation into the endometrial assembloids (Fig.6A). By Day 9, we observed that the blastoids could grow within the endometrial assembloids and interact with them (Fig.6B). The co-cultured blastoids displayed normal tri-lineage differentiation, characterized by markers for the epiblast (OCT4), hypoblast (GATA6), and trophoblast (KRT18) (Fig.6C). Additionally, we conducted a comparative analysis of the survival rates of blastoids and their interaction rates with endometrial assembloids across the CTRL, SEC and WOI assembloids. Remarkably, the survival and interaction rates of blastoids in the WOI endometrial assembloids were significantly higher compared to the CTRL and SEC groups, with survival rates of 66%, 19%, and 28%, and interaction rates of 90%, 47%, and 53% respectively, indicating the promising potential for embryo implantation (Fig.6D∼E). This demonstrates that we have provided an innovative and supportive model suitable for embryo implantation.

Revised version: To validate WOI assembloids’ ability to recapitulate embryo implantation (a key biological process of receptive endometrium), we established an assembloid-based co-culture system. Given the rarity and ethical constraints associated with human embryos, we employed blastoids (corresponding to the human embryo at 6 days post-fertilization, referred to as “Day 6”) for implantation into the endometrial assembloids (Fig.6A). By Day 9, blastoids were observed to survive and integrate within the endometrial assembloids (Fig.6B). Importantly, co-cultured blastoids retained the capacity for normal tri-lineage differentiation, as evidenced by the appropriate expression of epiblast (OCT4), hypoblast (GATA6), and trophoblast (KRT18) lineage markers (Fig.6C). We next quantified blastoids survival rates and their rates of interaction with endometrial assembloids across the CTRL, SEC, and WOI groups. Strikingly, blastoids co-cultured with WOI assembloids exhibited significantly higher survival and interaction rates compared to CTRL and SEC assembloids, with survival rates of 66% (WOI), 19% (CTRL), and 28% (SEC), and interaction rates of 90% (WOI), 47% (CTRL), and 53% (SEC) respectively (Fig.6D∼E). This demonstrates that WOI assembloids possess a markedly enhanced capacity to support blastoids survival and integration, highlighting their functional relevance to the receptive endometrium.

The receptive endometrial assembloids possess the potential for embryo implantation.

Added in the revised version: (A) Diagram illustrated the co-culture model of endometrial assembloids with blastoids (the blastoid stage corresponds to a 6-day post-fertilization human embryo, referred to as Day 6 here). mIVC1: modified In Vitro Culture Medium 1, mIVC2: modified In Vitro Culture Medium 2. (B) Bright-field images of the co-culture of CTRL, SEC, and WOI assembloids with blastoids (Day 9) (yellow arrows indicate the blastoids). Scale bar = 100μm. (C) Whole-mount fluorescence staining of Day 9 co-cultured embryoid and assembloid. OCT4 indicates the epiblast, GATA6 indicates the hypoblast, and KRT18 indicates the trophoblast. Scale bar = 40μm and 20μm (the rightmost image). (D) Comparison of the survival rates of Day 9 embryoid in CTRL, SEC, and WOI assembloids. (E) Comparison of the interaction ratios between Day 9 embryoid and endometrial assembloids in the CTRL, SEC, and WOI groups. *P < 0.05, **P < 0.005.

Removed in the revised version: (A) Diagram illustrated the co-culture model of endometrial assembloids with blastoids (the blastoid stage corresponds to a 6-day post-fertilization human embryo, referred to as Day 6 here).

Removed in the revised version: (B) Bright-field images of the co-culture of CTRL, SEC, and WOI assembloids with blastoids (Day 9) (yellow arrows indicate the blastoids). Scale bar = 100μm.

Removed in the revised version: (C) Whole-mount fluorescence staining of Day 9 co-cultured embryoid and assembloid. OCT4 indicates the epiblast, GATA6 indicates the hypoblast, and KRT18 indicates the trophoblast. Scale bar = 40μm and 20μm (the rightmost image).

Removed in the revised version: (D) Comparison of the survival rates of Day 9 embryoid in CTRL, SEC, and WOI assembloids.

Removed in the revised version: (E) Comparison of the interaction ratios between Day 9 embryoid and endometrial assembloids in the CTRL, SEC, and WOI groups. *P < 0.05, **P < 0.005.

In summary, WOI endometrial assembloids not only exhibited the typical structural and molecular features of the implantation window but also demonstrated significant potential for embryo implantation.

Discussion

Earlier version: In our study, we constructed the WOI endometrial assembloids, and observed the remarkable resemblance in structure and function to the in vivo endometrium. The assembloids consist of three primary types of endometrial cells, specifically epithelial, stromal, and immune cells, with epithelial cells assembling into glands and surrounded by immune cells and stromal cells. This similarity in cellular composition and tissue architecture to the endometrium lays the foundation for simulating receptive endometrium.

Revised version: In this study, we constructed the WOI endometrial assembloids, and observed the remarkable resemblance in structure and function to the in vivo endometrium. The assembloids consist of three primary types of endometrial cells, specifically epithelial, stromal, and immune cells, with epithelial cells assembling into glands surrounded by immune cells and stromal cells. This conservation of cellular composition and tissue architecture provides a foundational basis for mimicking the receptive endometrium.

Earlier version: Although previous studies treated endometrial assembloid with E2, P4, and cAMP to induce a transition to the secretory phase, the methodology to accurately simulate the implantation window or mid-secretory endometrium remains elusive and warrants further investigation. Previous research suggested that placental signals could further promote the differentiation of endometrial assembloids (2), and PRL, hCG and HPL have been implicated in processes such as decidualization, implantation, immunoregulation, and angiogenesis (21, 22, 24). Detailly, PRL, synthesized by the adenohypophysis, endometrium, and myometrium, plays a vital role in implantation, immunoregulation, and angiogenesis, with the secretory endometrium producing PRL in response to MPA and E2, leading to ciliated cell formation and stromal cell decidualization (2, 21). HCG, secreted by trophoblasts in early pregnancy, influences decidual cells (22) and improves endometrial thickness and receptivity(23). The introduction of hCG has been shown to elevate the expression of critical factors associated with endometrial receptivity, such as endocytosis proteins, hypoxia-inducible factor 1 (HIF1), chemokines, and glycodelin (53). HPL contributes to the development and function of uterine glands(24). Consequently, the additional supplementation of PRL, hCG and HPL in our system augmented hormone responsiveness and receptivity, while promoting cell differentiation, ultimately yielding a model more representative of the receptive endometrium.

Revised version: While previous studies have induced secretory phase transition in endometrial assembloids using E2, P4, and cAMP, methodologies for accurately simulating the WOI or mid-secretory endometrium remain inadequate and require further refinement. Prior research has indicated that placental signals can enhance endometrial assembloid differentiation (2), with PRL, hCG and HPL implicated in decidualization, implantation, immunoregulation, and angiogenesis (21, 22, 24). Specifically, PRL, synthesized by the adenohypophysis, endometrium, and myometrium, plays a vital role in implantation, immunoregulation, and angiogenesis, with the secretory endometrium producing PRL in response to MPA and E2, leading to ciliated cell formation and stromal cell decidualization (2, 21). HCG, secreted by trophoblasts in early pregnancy, modulates decidual cells (22) and improves endometrial thickness and receptivity(23). The introduction of hCG has been shown to upregulate key receptivity-associated factors, such as endocytosis proteins, hypoxia-inducible factor 1 (HIF1), chemokines, and glycodelin (52). HPL contributes to the development and function of uterine glands(24). Consequently, supplementary PRL, hCG and HPL in our system augmented hormone responsiveness and receptivity, promoted cellular differentiation, and ultimately generated a model more representative of the receptive endometrium.

Earlier version: Here, the WOI assembloid exhibited the characteristic ultra-structures, such as cilia. During the human menstrual cycle, motile cilia are present in the epithelium of the endometrium(1). These hair-like organelles extend from the cell surface and beat rhythmically, facilitating cell and tissue movement while driving fluid transport across the epithelium (54). During the decidualization of the endometrium, the number and length of cilia increase, a process driven by the two primary regulatory factors for embryo implantation: estrogen and progesterone (15, 55, 56). However, with aging, the expression of cilia-related genes in the endometrium is downregulated(57). In patients with recurrent implantation failure (RIF), ciliary defects in the endometrium are observed, leading to abnormal decidualization and repeated implantation failure(56). Thus, cilia play a crucial role in endometrial decidualization and embryo implantation. Consistently, the WOI assembloid exhibited increased ciliated epithelium, along with enhanced assembly and motility of cilia.

Revised version: Notably, WOI assembloids exhibited the characteristic ultrastructures, such as cilia. Motile cilia are present in endometrial epithelium throughout the human menstrual cycle (1). These hair-like organelles extend from the cell surface and beat rhythmically, facilitating cell and tissue movement while driving fluid transport across the epithelium (53). During endometrial decidualization, the number and length of cilia increase, a process driven by the two primary regulatory factors for embryo implantation: estrogen and progesterone (15, 54, 55). Throughout the menstrual cycle, ciliated and unciliated epithelia undergo mutual transformation from the secretory phase (or mid-secretory phase) to the menstrual phase, and then to the proliferative phase. Ciliated cell abundance peaks in the early-to-mid-secretory endometrium (1), with the mid-secretory phase exhibiting a higher proportion of ciliated cells as opposed to the early and late-secretory phases (25). However, with aging, the expression of cilia-related genes in the endometrium is downregulated(56). Recurrent implantation failure (RIF) patients often present with endometrial ciliary defects, leading to impaired decidualization and repeated implantation failure (55). Thus, cilia play a crucial role in endometrial decidualization and embryo implantation. Consistently, the WOI assembloid exhibited increased ciliated epithelium, along with enhanced assembly and motility of cilia.

Earlier version: Furthermore, the WOI assembloid demonstrated energy and lipid metabolism patterns resembling those of the in vivo receptive endometrium. Enhanced energy metabolism involves monocarboxylic acid metabolism and mitochondrial oxidative phosphorylation. Monocarboxylic acids, exemplified by lactate, pyruvate, and ketone bodies, are vital metabolites in most mammalian cells. During the implantation window, elevated lactate levels mobilize endometrial monocarboxylic acid metabolism and function as a pregnancy-related signal, stimulating secretion in the epithelium. Subsequently, ATP, the primary product of energy metabolism, induces neighboring epithelial cells to release IL8, promoting decidualization of stromal cells (42). It was also observed that WOI assembloids indeed possessed larger and functionally active mitochondria, producing much more ATP and IL8 than CTRL and SEC assembloids. Lipid metabolism, responsible for energy storage, signal transduction, cell proliferation, apoptosis, and membrane trafficking, plays a crucial role in endometrial receptivity and implantation, although the precise mechanisms are not fully understood (58–60). Thus, the WOI assembloids possessed metabolic characteristics of in vivo implantation window.

Revised version: Furthermore, WOI assembloids demonstrated energy and lipid metabolism patterns resembling those of the in vivo receptive endometrium. Enhanced energy metabolism involves monocarboxylic acid metabolism and mitochondrial oxidative phosphorylation. Monocarboxylic acids, exemplified by lactate, pyruvate, and ketone bodies, are vital metabolites in most mammalian cells. During the implantation window, elevated lactate levels mobilize endometrial monocarboxylic acid metabolism and function as a pregnancy-related signal, stimulating secretion in the epithelium. Subsequently, ATP, the primary product of energy metabolism, induces neighboring epithelial cells to release IL8, promoting decidualization of stromal cells (42). Concordantly, our WOI assembloids indeed possessed larger and functionally active mitochondria, and produced much more ATP and IL8 than CTRL and SEC assembloids. Lipid metabolism, responsible for energy storage, signal transduction, cell proliferation, apoptosis, and membrane trafficking, plays a crucial role in endometrial receptivity and implantation, although its precise mechanisms remain incompletely understood (57–59). Thus, WOI assembloids possessed metabolic characteristics of in vivo implantation window.

Earlier version: However, our WOI endometrial assembloids also exhibit some limitations. It is undeniable that the assembloids cannot perfectly replicate the in vivo endometrium, which comprises functional and basal layers with a greater abundance of cell subtypes, under superior regulation by systemic hormones and nutrients. Stromal cells and immune cells are difficult to pass down stably and their proportion is lower than that in the in vivo endometrium. We are looking forward to combining stem cell induction, 3D printing technology, and microfluidic systems to alter the culture environment in order to optimize the existing model.

Revised version: However, our WOI endometrial assembloids also exhibit some limitations. It is undeniable that the assembloids cannot perfectly replicate the in vivo endometrium, which comprises functional and basal layers with a greater abundance of cell subtypes, under superior regulation by hypothalamic-pituitary-ovarian (HPO) axis. Specifically, stromal and immune cells are challenging to stably passage, and their proportion is lower than in the in vivo endometrium. While the in vivo peri-implantation period exhibits intricate gene expression dynamics driven by systemic regulation, our models only partially recapitulate these changes, primarily in mitochondria-and cilia-associated genes. Nevertheless, to some extent, these WOI assembloids possess receptivity characteristics and can be utilized for investigating receptivity-related scientific questions or conducting in vitro drug screening. Further refinements are required to fully simulate the dynamic endometrial gene expression patterns across all menstrual cycle stages. We are looking forward to integrating stem cell induction, 3D printing, and microfluidic systems to modify the culture environment.

Earlier version: In summary, we developed a receptive endometrial assembloid that mimics the features of in vivo implantation-window endometrium (Fig. 7). With typical ultra-structures of the implantation window, including pinopodes and cilia, it not only demonstrated hormonal responsiveness and characteristic glycogen secretion functions, but also recapitulated the processes such as decidualization, metabolic changes and EMT, while retaining the potential for embryo implantation. This receptive endometrial assembloid serves as a platform to investigate peri-implantation endometrial physiology and pathology, maternal-fetal interactions, with potential practical applications and clinical translation.

Revised version: In summary, we developed a receptive endometrial assembloid that mimics in vivo implantation-window endometrium (Fig. 7). This model exhibits WOI-characteristic ultrastructures, including pinopodes and cilia, displays hormonal responsiveness and characteristic glycogen-secreting functions, recapitulates the processes such as decidualization, metabolic changes and EMT, and retains embryo implantation potential. It serves as a valuable platform to investigate peri-implantation endometrial physiology and pathology, maternal-fetal interactions, with potential practical applications and clinical translation.

Schematic diagram displaying the establishment and validation of receptive endometrial assembloids, and summarizing the characteristic biological events of implantation-window endometrium.

Materials and Methods

Establishment of endometrial assembloids

All experiments involving human subjects followed medical ethical principles and the Declaration of Helsinki, and were approved by the Ethics Committee of Qilu Hospital of Shandong University. Informed consents were obtained from patients. Females of reproductive age who received hysterectomy for benign diseases were selected for this study. Clinical information of patients providing endometrial tissue was listed in Dataset S1. Experimental procedures are presented in detail in the Supporting Information.

Hormone treatment of endometrial assembloids

Earlier version: The hormone regimen for inducing endometrial secretory phase, as described by Thomas E. Spencer et al (9), consisted of estradiol for 2 days followed by a combination of estradiol (Sigma E2758), medroxyprogesterone acetate (Selleck S2567) and N6,2′-O-dibutyryladenosine 3′,5′-cyclic monophosphate sodium salt (cAMP) (Sigma D0627) for 6 days.

Revised version: Hormonal treatment was initiated following the assembly of the endometrial assembloids (about 7-day growth period). The hormone regimen for inducing endometrial secretory phase, as described by Thomas E. Spencer et al (9), consisted of estradiol for 2 days followed by a combination of estradiol (Sigma E2758), medroxyprogesterone acetate (Selleck S2567) and N6,2′-O-dibutyryladenosine 3′,5′-cyclic monophosphate sodium salt (cAMP) (Sigma D0627) for 6 days.

Earlier version: To simulate the endometrium of implantation window, we developed a model by incorporating various pregnancy-related hormones. The hormone regimen included estradiol for the first two days, followed by a combination of estradiol, medroxyprogesterone acetate, cAMP, Human Chorionic Gonadotropin (HCG) (Livzon Pharmaceutical Group Inc), Human Placental Lactogen (HPL) (R&D Systems 5757-PL), and prolactin (Peprotech 100-07) for 6 days. The CTRL group was cultured in ExM without the addition of hormones. (Fig. 1A) (Table S2)

Revised version: To simulate the endometrium of implantation window, we developed a model by incorporating various pregnancy-related hormones. The hormone regimen included estradiol for the first two days, followed by a combination of estradiol, medroxyprogesterone acetate, cAMP, Human Chorionic Gonadotropin (HCG) (Livzon Pharmaceutical Group Inc), Human Placental Lactogen (HPL) (R&D Systems 5757-PL), and prolactin (Peprotech 100-07) for 6 days. The CTRL group was cultured in ExM without hormone supplementation and subjected to parallel culture for 8 days along with the two aforementioned groups. (Fig. 1A) (Table S2)